University of Michigan - ROB 101 Computational Linear Algebra

# Homework 1: How to use Julia as a Calculator
### Due: 9/10 at 9 pm Eastern Time

#### Pupose:  Use Julia Programming to solve mathematical problems
- Skills:
    - Writing and reading code
    - Simplifying tasks with computer programs
- Knowledge to be gained:
    - Basic math commands
    - Types of variables in Julia
    - Declaring and plotting functions
    
#### Task:
Welcome to your first jupyter notebook assignment!  To complete this assignment, fill in your code as instructed.  

Anywhere you see a "throw(ErrorException())", delete that line, and replace it with your own code.  Once you submit your assignment, it will be auto-graded using the the code you added.  These auto-grader cells will sometimes be visible to you, as a way for you to check your code before submitting the assignment.  

Sometimes, you will have to complete tasks which don't have the "throw(ErrorException())".  In that case, follow the instructions in the cell, and email your teachers if you have questions.  Some cells do not have any tasks for you, but are either there for you to play with, or necessary structures for some of the programming.  


## Julia can be even more powerful than a graphing calculator if you know how to use it!


The Basics:
Arithmetic Operations!

```julia
#in a code cell, julia will compute and return operations you do inside it
4+7.62
#here is some basic addition
```

```julia
#this works the same way for addition(+), subtraction(-), multiplication(*), division(/) and exponentials(^)

4+2/3-1*4

#Julia also has a backwards division operator (\) to divide the left by the right
#and a (%) modulo operator, which returns the remainder from a division operation
```

**Notice: Julia will do arithmetic operations in an order that follows PEMDAS**
1. Parenthesis
2. Exponentials (right to left)
3. Multiplication and Division (left to right)
4. Addition and Subrtraction(left to right)

Order of operations is extremely important, and can cause errors in your code so PAY ATTENTION!

### Problem 1.  Place two sets of parentheses in the below expression to make x equal to 50

```julia
#place your parenthesis here
x = 4+ 5/2+3 + 2^2 +1

```

<!-- grade_id: problem1test, points: 2, grade: true, solution: false -->
```julia
#This is an autograder cell
@assert x == 50
```

**Julia also has many built in functions:**

- round() #rounds to nearest integer
- ceil() #rounds up
- floor() #rounds down
- trunc() #chops off anything beyond decimal point
- abs() #absolute value
- sin(), cos(), tan()  and asin(), acos(), atan()  #many more trigonometric functions
- sqrt() #square root
- exp()  #natural exponential e^x
- log() #natural logarithm

and so many more: https://docs.julialang.org/en/v1/manual/mathematical-operations/

Note: These operations will only work on numbers! In each cell, hit shift return

```julia
sin(pi)  #It assumes radians and not degrees
```

```julia
sqrt(144)  #yep, the usual square root of a number
```

```julia
abs("hello")  #be prepared for an error!
```

Note: This doesn't work because "hello" is the wrong type of data.

**Every piece of data in Julia has a "Type"**

What is a data type? I think you would be comfortable saying that numbers are different from letters. So, they should be treated differently when you work with them. Similary, a single letter is different than a sentence. 

So far, so good? 

Here is where programming languages are different than calculators: you may not have to think about numbers such as 1, 2, 3, 4, .... differently than 0.7, sqrt(2), pi, but Julia does have to handle them differently. 

Why? Just to make your life miserable? Not really. 

Computers have a limited amount of memory (zeros and ones) available to represent a number. If the computer knows that the number does not require a decimal point (as in -7, 0, 1, 45, etc,) then it can store it very compactly and hence store enormously large numbers of the non-decimal TYPE. If the number requires a decimal to express it properly (such as pi, sqrt(2), 0.6) then Julia will store it in a different format. In Julia, formats are determined by TYPES.

We will mostly use 4 of these TYPES:
- Int64:  Integers (numbers without decimal points in them)
- Float64: Floating point numbers (numbers that require decimal points)
- Char:  Any character you can type on a keyboard, or fancy characters like 复, 制, ¤, Ð, or ϐ (declared inside ' ')
- String: A bunch of characters in a line.  (This could be a word, or a sentence, or even a number that is in " ")

You don't usually have to worry about types because Julia will automatically choose a type that works best.  But, should you get a MethodError like the one above, check which type you are using.

If you want to know more about types in Julia: https://docs.julialang.org/en/v1/manual/types/

**Solving and graphing your own functions**

Julia can also help solve any equation that you input.  There are two main ways to declare a function in Julia.

Let's practice with the basic linear function:  
                                                y = 2x 

The first one is the most simple:

```julia
f(x) = 2x

#now that we have it declared, we can plug in values of x
f(3)
```

Now, we want to plot this function, to see where it's solutions are.

```julia
#These lines give Julia the ability to plot functions.  This may take a while, but it is necessary.
# When you see [*], that means Julia is working hard to do what you told it. Your commands may take just a few 
# lines, but Julia is probably dealing with multiple megabytes of files and code, all of which is hidden from us. 
using Plots
gr()
```

```julia
#now, we can use the plot() function
plot(f, -5, 5)  #f will be plotted over an x range of -5 to 5
```

If we wanted to change an element of the graph, we need to use the function **plot!()**

The exclamation point indicates that the function will modify the existing plot, instead of replacing it.  Let's make the graph to scale with an aspect_ratio command, and label the axes.  You can also add an additional line to the plot.  Let's include the x-axis.

```julia
plot!(aspect_ratio = 1, xlabel = "x", ylabel = "y")
g(x) = 0
plot!(g, -15,15)
```

We can now find that y = 2x intersects y = 0 at (0,0).  
### Problem 2. Practice with one of the systems in your HW1 Assignment

                                                    x − 2y = −2
                                                     x + y = 4
                                                
                                                
                                                
Step 1:  Rearrange Both equations so that y is isolated on the left side:

Equation 1: y1 = ____________
Equation 2: y2 = ____________

<!-- grade_id: problem2, points: 0, grade: false, solution: true -->
```julia
#now declare both of those equations as code functions y1(x) and y2(x)

# your code here
throw(ErrorException())

#Plot them on the same graph.  This shows another way to plot multiple things at once
plot([y1,y2], -5, 5)
```

<!-- grade_id: problem2test, points: 2, grade: true, solution: false -->
```julia
#Grader Cell
@assert y1(0) == 1
@assert y2(0) == 4
```

Now that you have plotted the system, you can find the solution as the intersection of the lines.  Pretty neat!

### Problem 3. Creating Arrays
### Create a row vector called rowVect that contains the numbers 1 through 10 using the collect function, see the example code below

```julia
# Run me, don't change me 
#output should be
#1 2 3 4 5 6 7 8 9 10
# Example Code. 
rowVect = collect(1:10)'

```

<!-- grade_id: cell-2d4d8dec7be601ab, points: 0, grade: true, solution: false -->
```julia
#autograder cell
#Just checking if you ran the code cell! 
```

### a) Create a column vector called colVect that contains the numbers 11 through 20

<!-- grade_id: cell-d374a9606cd3edc8, points: 0, grade: false, solution: true -->
```julia
#output should be
#11 
#12
#13
#14
#15
#16
#17
#18
#19
#20

# colVect = ??
# your code here
throw(ErrorException())
```

<!-- grade_id: cell-7df9c508380d7636, points: 2, grade: true, solution: false -->
```julia
#autograder cell
```

**Note: Other ways to build vectors**
* Range notation n:k:m will start at n and take steps of size k to reach m.  Use collect(n:k:m) to make it into a vector
* A[:] converts a matrix A to a vector
* diag() returns the diagonal elements of a matrix as a vector

### Problem 4.
Using Indexing. we can extract cvalues from vectors by indexing variables, let's see how!


```julia
# What is the value of the 5th member of the colVect vector we just created, can you verify with this code?
@show value = colVect[5]
```


### a) Use indexing [] to return the "panda".  Keep in mind that Julia indexes refer to the location in a vector.

<!-- grade_id: cell-0be31681aba733b4, points: 0, grade: false, solution: true -->
```julia
animalsVector = ["lemur", "elephant", "tiger", "panda", "zebra", "cuttlefish"]
# myPanda = ??
# your code here
throw(ErrorException())
```

<!-- grade_id: cell-51e5657dac47588e, points: 2, grade: true, solution: false -->
```julia
#autograder cell
@assert myPanda == "panda"
```

### Problem 5. Special Matrix commands

Let's start with an example code. Let's create a matrix of zeros with 2 rows and 3 columns

```julia
Example_zeroMat = zeros(2,3)
```

### a) Zeros: Create a matrix called zeroMat of zeros with 4 rows and 7 columns

<!-- grade_id: cell-9ca6613d99726a6b, points: 0, grade: false, solution: true -->
```julia
# your code here 
# zeroMat = ??
# your code here
throw(ErrorException())
```

<!-- grade_id: cell-d124ce482005f752, points: 2, grade: true, solution: false -->
```julia
#autograder cell
# Just checking if you succesfully ran the code! 
```

### Learning from the special matrix command zeros can you create a matrix of random numbers?
Hint: Use the statement ```randn``` (equivalent of zeros but for creating matrices of random numbers) 
For example see the code in the fourth cell after this, it creates a 1x1 Matrix of a random number!

### b) Create a matrix A of random numbers from a standard normal distribution(randn) with dimensions (17,8) 

<!-- grade_id: cell-7aef9ce25edc0260, points: 0, grade: false, solution: true -->
```julia
# A = randn(?,?)
# your code here
throw(ErrorException())
```

<!-- grade_id: cell-8cd8412de17d5b8f, points: 2, grade: true, solution: false -->
```julia
#autograder cell
```

### A note about random seeds

Random numbers, in Julia, are not truly "random". They are instead generated from a seed. In future homeworks, we will set the random seed manually, so that we are able to automatically grade your code output, even when the code inputs are ostensibly random. This property is also good for you - it would be difficult, to say the least, to diagnose issues in output if your input changes every time you run the cell! Here is a demonstration: the following cell, that generates a random number based on a seed, will generate the random number <code>-0.5560268761463861</code>

```julia
using Random 
Random.seed!(42)
random_num = randn(1,1)
```

Whoah! As can be seen, the random number generated is exactly the same. CAUTION: IT IS CRUCIAL TO ALWAYS MAKE SURE CELLS CONTAINING RANDOM SEEDS IN HOMEWORKS ARE RUN BEFORE RUNNING CELLS THAT GENERATE RANDOM NUMBERS. Humans may be smart, but autograders are dumb: we know that you generated a random number that we did not predict, but the autograder just sees an unexpected output, and assumes that your code was incorrect.

**Note:  Other ways to build matrices**
* use hcat and vcat to concatenate vectors together horizontally or vertically
* use ones() or zeros(m, n) to create a m by n matrix of either ones or zeros
* use Diagonal(M) to create a diagonal matrix from a vector M
* Random functions
    * rand(m, n) makes an m by n matrix of random numbers from a uniform distribution 0-1
    * randn(m, n) m by n matrix from a standard normal distribution
    * rand(X, m, n) m by n matrix of random numbers in an array X
* use reshape(X, m, n) to turn a vector X into a matrix size m by n
* use repeat(X, m ,n) to repeat matrix X m times in the row direction and n times in the column direction