University of Michigan - ROB 101 Computational Linear Algebra

# Julia Lab 01: How to use Julia as a Calculator

#### Pupose:  Use Julia Programming to solve mathematical problems
- Skills:
    - Writing and reading code
    - Simplifying tasks with computer programs
- Knowledge to be gained:
    - Basic math commands
    - Types of variables in Julia
    - Declaring and plotting functions
    
#### Task:
Welcome to your first jupyter notebook experience! We will complete this assignmnet together.  

Anywhere you see a "throw(ErrorException())", delete that line, and replace it with your own code.

Sometimes, you will have to complete tasks which don't have the "throw(ErrorException())".  In that case, follow the instructions in the cell, and email your teachers if you have questions.  Some cells do not have any tasks for you, but are either there for you to play with, or necessary structures for some of the programming.  

### Enable Line Numbering

In labs and on homework assignments, you may come accross instructions which tell you to alter certain lines of your code. You will also need to debug your code at times based on the errors reported by Julia, which will tell you the line number at which your program failed. In order to easily alter your code, enable line numbering by clicking View > Show Line Numbers in the Jupyter Lab/Notebook toolbar. 

## Julia can be even more powerful than a graphing calculator if you know how to use it!


The Basics:
Arithmetic Operations!

```julia
#in a code cell, julia will compute and return operations you do inside it
4+7.62
#here is some basic addition
```

```julia
#this works the same way for addition(+), subtraction(-), multiplication(*), division(/) and exponentials(^)

4+2/3-1*4

#Julia also has a backwards division operator (\) to divide the left by the right
#and a (%) modulo operator, which returns the remainder from a division operation
```

**Notice: Julia will do arithmetic operations in an order that follows PEMDAS**
1. Parenthesis
2. Exponentials (right to left)
3. Multiplication and Division (left to right)
4. Addition and Subrtraction(left to right)

Order of operations is extremely important, and can cause errors in your code so PAY ATTENTION!

### Problem 1.  Place two sets of parentheses in the below expression to make x equal to 49

```julia
#place your parenthesis here
x = 3 + 5 * 3 + 4 / 2 * 4 + 5
x = 49
```

<!-- grade_id: problem1test, points: 2, grade: true, solution: false -->
```julia
#This is your first introduction to an autograder cell
@assert x == 49
```

<!-- grade_id: problem1test_fail, points: 0, grade: true, solution: false -->
```julia
#This is your first introduction to an autograder cell that FAILS!
#Notice that the assert statement is wrong... Julia won't be happy!
@assert x == 5
```

**Julia also has many built in functions:**

- round() #rounds to nearest integer
- ceil() #rounds up
- floor() #rounds down
- trunc() #chops off anything beyond decimal point
- abs() #absolute value
- sin(), cos(), tan()  and asin(), acos(), atan()  #many more trigonometric functions
- sqrt() #square root
- exp()  #natural exponential e^x
- log() #natural logarithm

and so many more: https://docs.julialang.org/en/v1/manual/mathematical-operations/

Remember, Julia is a mathematical programming language - it can do many of the mathematical programming operations that you are used to.

Note: These operations will only work on numbers! In each cell, hit shift return

```julia
cos(2*pi)  #It assumes radians and not degrees
```

```julia
exp(log(2))  #looks correct - remember e ^ log(x) = x
```

```julia
ceil(1.5)    #note that ceil and floor always round "predictably"
```

```julia
floor(1.5)    #for instance, floor always returns a number less than or equa to the input argument
```

```julia
abs("hello")  #be prepared for an error!
```

Note: This doesn't work because "hello" is the wrong type of data. We can see this evidenced in the error message: it tells us "no method matching abs(::String) because we are trying to use a function names "abs" that takes an argument of type "String" (a word, or a bunch of letters/characters). Of course, we know that you obviously can not take the absolute value of a word or letter, but Julia does not. Thus, when creating a function, checking the type of the argument (the thing we pass in to the function) helps to make sure the function works correctly.

Where did the error occur? Let us look at a slightly different cell:

```julia
x = abs(sin(2))
abs("hello")  #be prepared for an error!
x = x + 2
```

We have three lines in the above cell (by the way, this would be a good time to enable line numbering in View > Show Line Numbers). On the first line, we assign the variable x the absolute value of the sin of 2. On the second line, we take the absolute value of the string "hello". FInally, on the third line, we increment x by 2. When we run the above cell, we see that Julia is angry with us, as it tells us <code>MethodError: no method matching abs(::String)MethodError: no method matching abs(::String)</code>. But, where did this error occur? Julia tells us this if we look under the <code>Stacktrace</code> section of the error message. The top-level scope (AKA where train came off the tracks) is <code>In[some number]:2</code>. Or, in some numbered cell (the number of the cell is determined when you run the cell: if this is the first cell run in the notebook, it will be cell 1) on line 2. This is what we expected - after all, you can't take the absolute value of a string, how can a word have a sign!

**Every piece of data in Julia has a "Type"**

What is a data type? I think you would be comfortable saying that numbers are different from letters. So, they should be treated differently when you work with them. Similary, a single letter is different than a sentence. 

So far, so good? 

Here is where programming languages are different than calculators: you may not have to think about numbers such as 1, 2, 3, 4, .... differently than 0.7, sqrt(2), pi, but Julia does have to handle them differently. 

Why? Just to make your life miserable? Not really. It's actually to make your life easier! 

Computers have a limited amount of memory (zeros and ones) available to represent a number. If the computer knows that the number does not require a decimal point (as in -7, 0, 1, 45, etc,) then it can store it very compactly and hence store enormously large numbers of the non-decimal TYPE. If the number requires a decimal to express it properly (such as pi, sqrt(2), 0.6) then Julia will store it in a different format. Thus, Julia is able to handle the memory needs of our code without our help, because we essentially tell Julia in advance the memory needs of our code. In Julia, formats are determined by TYPES.

We will mostly use 4 of these TYPES:
- Int64:  Integers (numbers without decimal points in them). The 64 refers to the number of 'bits', 0s and 1s, that will represent the number.
- Float64: Floating point numbers (numbers that require decimal points)
- Char:  Any character you can type on a keyboard, or fancy characters like 复, 制, ¤, Ð, or ϐ (declared inside ' ')
- String: A bunch of characters in a line.  (This could be a word, or a sentence, or even a number that is in " ")

You don't usually have to worry about types because Julia will automatically choose a type that works best.  But, should you get a MethodError like the one above, check which type you are using.

If you want to know more about types in Julia: https://docs.julialang.org/en/v1/manual/types/

**Solving and graphing your own functions**

Julia can also help solve any equation that you input.  There are two main ways to declare a function in Julia.

Let's practice with the basic linear function:  
                                                $$y = 5x + 2$$

The first one is the most simple:

```julia
f(x) = 5x + 2

#now that we have it declared, we can plug in values of x
f(3)

#See that we return 17: (5 * 3) + 2
```

Now, we want to plot this function, to see where it's solutions are. To do this, we will use the "Plots" package. Julia is a fairly lightweight programming language - it does not naturally do more than the bare essentials. We then add functionalities that we need by "importing", or downloading, packages of code that add functionalities which do useful things, like we are about to see.

```julia
#These lines give Julia the ability to plot functions.  This may take a while, but it is necessary.
# When you see [*], that means Julia is working hard to do what you told it. Your commands may take just a few 
# lines, but Julia is probably dealing with multiple megabytes of files and code, all of which is hidden from us. 

import Pkg
Pkg.add("Plots")
using Plots

gr()
```

```julia
#now, we can use the plot() function
plot(f, -5, 5)  #f will be plotted over an x range of -5 to 5
```

If we wanted to change an element of the graph, we need to use the function **plot!()**

The exclamation point indicates that the function will modify the existing plot, instead of replacing it.  Let's make the graph to scale with an aspect_ratio command, and label the axes.  You can also add an additional line to the plot.  Let's include the x-axis.

```julia
plot!(aspect_ratio = 1, xlabel = "x", ylabel = "y")
g(x) = 0
plot!(g, -15,15)
```

We can now find that y = 2x intersects y = 0 at (0,0).  
### Problem 2. Finding the solution to a system of equations

                                                    2x − y = −1
                                                     4x + y = 4
                                                
                                                
                                                
Step 1:  Rearrange Both equations so that y is isolated on the left side:

Equation 1: y1 = ____________
Equation 2: y2 = ____________

<!-- grade_id: problem2, points: 0, grade: false, solution: true -->
```julia
#now declare both of those equations as code functions y1(x) and y2(x)

# your code here
throw(ErrorException())

#Plot them on the same graph.  This shows another way to plot multiple things at once
plot([y1,y2], -5, 5)
```

<!-- grade_id: problem2test, points: 2, grade: true, solution: false -->
```julia
#Grader Cell
@assert y1(0) == 1
@assert y2(0) == 4
```

Now that you have plotted the system, you can find the solution as the intersection of the lines.  Pretty neat!

### Problem 3. Creating Arrays
### a)  Create a row vector called rowVect that contains the numbers 1 through 10

<!-- grade_id: problem3a, points: 0, grade: false, solution: true -->
```julia
#output should be
#1 2 3 4 5 6 7 8 9 10
# your code here
throw(ErrorException())
```

<!-- grade_id: problem3a_test, points: 2, grade: true, solution: false -->
```julia
#autograder cell
```

### b) Create a column vector called colVect that contains the numbers 11 through 20

<!-- grade_id: problem3b, points: 0, grade: false, solution: true -->
```julia
#output should be
#11 
#12
#13
#14
#15
#16
#17
#18
#19
#20

# your code here
throw(ErrorException())
```

<!-- grade_id: problem3b_test, points: 2, grade: true, solution: false -->
```julia
#autograder cell
```

**Note: Other ways to build vectors**
* Range notation n:k:m will start at n and take steps of size k to reach m.  Use collect(n:k:m) to make it into a vector
* A[:] converts a matrix A to a vector by concatenating the columns of the matrix A
* diag() returns the diagonal elements of a matrix as a vector

### Problem 4.
a) Use indexing [] to return the "panda".  Keep in mind that Julia indexes refer to the location in a vector.

<!-- grade_id: problem4a, points: 0, grade: false, solution: true -->
```julia
animalsVector = ["lemur", "elephant", "tiger", "panda", "zebra", "cuttlefish"]
myPanda = 
# your code here
throw(ErrorException())
```

<!-- grade_id: problem4a_test, points: 2, grade: true, solution: false -->
```julia
#autograder cell
@assert myPanda == "panda"
```

b) Use indexing to return the "lemur". (Hint: what index number do julia vectors start from?)

<!-- grade_id: problem4b, points: 0, grade: false, solution: true -->
```julia
myLemur = 
# your code here
throw(ErrorException())
```

<!-- grade_id: problem4b_test, points: 2, grade: true, solution: false -->
```julia
#autograder cell
@assert myLemur == "lemur"
```

### Problem 5. Special Matrix commands
### a) Create a matrix called onesMat of ones with 3 rows and 4 columns

<!-- grade_id: problem5a, points: 0, grade: false, solution: true -->
```julia
#Hint: see note below for a shortcut method
# your code here
throw(ErrorException())
```

<!-- grade_id: problem5a_test, points: 2, grade: true, solution: false -->
```julia
#autograder cell
```

### b) Create a matrix A of random numbers from a standard normal distribution with dimmensions (17,8) as shown in the autograder cell.

<!-- grade_id: problem5b, points: 0, grade: false, solution: true -->
```julia
#Hint:  see note below 
# your code here
throw(ErrorException())
```

<!-- grade_id: problem5b_test, points: 2, grade: true, solution: false -->
```julia
#autograder cell
#This is how all of your answers in the homework will be graded:
#An autograder cell, like this one, will check certain elements of your answer against what they should be.
#Normally, of course, the elements that are checked will be hidden. 
(nrow, ncol) = size(A) #do not change this line!
```

**Note:  Other ways to build matrices**
* use hcat and vcat to concatenate vectors together horizontally or vertically
* use ones() or zeros(m, n) to create a m by n matrix of either ones or zeros
* use Diagonal(M) to create a diagonal matrix from a vector M
* Random functions
    * rand(m, n) makes an m by n matrix of random numbers from a uniform distribution 0-1
    * randn(m, n) m by n matrix from a standard normal distribution
    * rand(X, m, n) m by n matrix of random numbers in an array X
* use reshape(X, m, n) to turn a vector X into a matrix size m by n
* use repeat(X, m ,n) to repeat matrix X m times in the row direction and n times in the column direction

### Gee Whiz, ain't that cool!
The below Epicycloid is a good reminder of what can be accomplished with Julia. Independently, our parametric equations for x and y are pretty simple, but when combined, they ca create something really interesting, as seen below. Play around with the values for a, b, and the upper bound of t - there are some beutiful patterns that can be achieved with the below functions.

Note: See how we have to call cos.(t)? You'll see kind of syntax  later when we start talking about matrix operations, but that dot tells julia to apply the function to each vector element individually, rather than the entire vector at once. Very helpful!

```julia
using Plots
t = 0 : 0.1 : 100
a = 4
b = 7
x = (a + b)cos.(t) - b*cos.(((a/b) + 1)t)
y = (a + b)sin.(t) - b*sin.(((a/b) + 1)t)
plot(x, y, aspect_ration=:equal)
```

```julia

```

```julia

```