University of Michigan - ROB 101 Computational Linear Algebra

# Homework 3: Vector and Matrix Multiplication
### Due: 9/26 at 9 PM Eastern

#### Purpose:  Learn ways to compute a matrix multiplication with programming
- Skills: 
    - Indexing for multiplication
    - Converting algorithm to code
    - Lower-Upper factorization without permutations
- Knowledge:
    - Learn how to use built in functions
    - Understand how those functions work behind the scenes

#### Task:  Complete the cells below as directed.

### The Result of a Matrix Multiplication:
$$[n\times k~~\text{matrix}] \cdot [k \times m~~\text{matrix}] = [n \times m~~\text{matrix}].$$

In other words:  In order for the matrix multiplication A times B to be possible, the number of columns in A must equal the number of rows in B.

\begin{equation}
   A \cdot B:= 
\left[\begin{array}{cccc}  a_1^{\rm row} \cdot b_1^{\rm col} & a_1^{\rm row} \cdot b_2^{\rm col} & \cdots & a_1^{\rm row} \cdot b_m^{\rm col} \\
a_2^{\rm row} \cdot b_1^{\rm col} & a_2^{\rm row} \cdot b_2^{\rm col} & \cdots & a_2^{\rm row} \cdot b_m^{\rm col} \\
\vdots & \vdots & \ddots & \vdots \\
a_n^{\rm row} \cdot b_1^{\rm col} & a_n^{\rm row} \cdot b_2^{\rm col} & \cdots & a_n^{\rm row} \cdot b_m^{\rm col}
\end{array}\right].
\end{equation}

### Problem 1.  Rows x Columns
Partition the following matrices, then compute the multiplication product using the formula above.

### a) $A*B$ where $A = \begin{bmatrix} 3&7&4\\1&4&3\end{bmatrix}$ and $B = \begin{bmatrix}8&3&9&1\\2&4&0&8\\0&7&1&2\end{bmatrix}$

<!-- grade_id: 1a, points: 0, grade: false, solution: true -->
```julia
#declare your partitions first
#leave the values as integers or the check below will return a false answer, even if you are right.
a1row = 
a2row = 
b1col = 
b2col = 
b3col = 
b4col =

#implement code below that computes the product using your partitioned matrices
#your answer should be a 2 x 4 matrix. Return that in a variable called answer1a
# your code here
throw(ErrorException())
```

<!-- grade_id: 1aTest, points: 1, grade: true, solution: false -->
```julia
#autograder cell
```

### b) $C*D$ where $C = \begin{bmatrix} 7&6&8&7&4\\1&8&2&4&3\\6&9&8&1&4\end{bmatrix}$ and $D = \begin{bmatrix}8&3\\0&4\\7&2\\3&8\\9&2\end{bmatrix}$

<!-- grade_id: 1b, points: 0, grade: false, solution: true -->
```julia
#Partition C 



#Partition D





#Compute product and return answer in a variable called answer1b
# your code here
throw(ErrorException())
```

<!-- grade_id: 1bTest, points: 1, grade: true, solution: false -->
```julia
#autograder cell
```

$$A\cdot B = \sum_{i=1}^{k} a_i^{\rm col} \cdot b_i^{\rm row},  $$
### Problem 2.  Summing over Columns and Rows
Partition the following matrices, then compute the product of matrix multiplication by summing over rows and columns.

### a.  $F*G$ where $F = \begin{bmatrix}5&1&3\\8&2&0\\1&7&6\\8&2&1\end{bmatrix}$ and $G = \begin{bmatrix}2&4\\8&1\\7&0\end{bmatrix}$

<!-- grade_id: 2a, points: 0, grade: false, solution: true -->
```julia
#partition F into columns



#partition G into rows



#compute the product by summing the products of the columns times rows.  Return your answer as answer2a
# your code here
throw(ErrorException())
```

<!-- grade_id: 2aTest, points: 1, grade: true, solution: false -->
```julia
#autograder cell
```

### b.  Using the same matrices, F and G, construct a FOR loop to compute the matrix multiplication $F*G$ like how we did above.
Remember this formula, and how it works:$$A\cdot B = \sum_{i=1}^{k} a_i^{\rm col} \cdot b_i^{\rm row},  $$

Hint: How many times do we have to multiply a column times a row?  What happens to all of the individual column times row matrices at the end?  Keep track of the sum in the for loop so that you don't have to manually sum at the end.

<!-- grade_id: 2b, points: 0, grade: false, solution: true -->
```julia
#Declare matrices F and G:
F = 
G = 
#Declare a matrix of zeros the size of the multiplication product
currentSum = 
#fill in the number of columns of F (or Rows of G)
n =

#the for loop
for i in 1:n
    #use indexing, and the iterator variable(i) to pick out the column and row,
    #then mulitply them together, and update the currentsum
    #In order to extract a row vector, remember to use the transpose operator (')
    
# your code here
throw(ErrorException())
    
end

answer2b = currentSum
```

<!-- grade_id: 2bTest, points: 2, grade: true, solution: false -->
```julia
#autograder cell
#You should get the same answer that you got in problem 2a
answer2b == answer2a
```

### The Trick:  Matrix Multiplication is Already Built into Julia
#### You can simply use the * operator to multiply compatible matrices
What?  I did all that for nothing!  No.  By completing the algorithms above, you have followed along the steps that Julia is taking internally. You now understand the inner workings of the computer!

Note: You still need to implement problems 1 and 2 as directed, otherwise you will receive no points on the assignment.

### Problem 3.  Compute $X*Y$ where $X = \begin{bmatrix}8&5&2&1\\6&1&3&2\end{bmatrix}$ and $Y = \begin{bmatrix}4&4&3\\7&6&5\\0&1&3\\9&8&2\end{bmatrix}$

<!-- grade_id: 3, points: 0, grade: false, solution: true -->
```julia
#Return your answer as answer3

# your code here
throw(ErrorException())
```

<!-- grade_id: 3test, points: 1, grade: true, solution: false -->
```julia
#autograder cell
```

<img src="https://i.postimg.cc/Y0HQcjMM/lu-final.jpg" width="700">


### Problem 4. LU Factorization Step by Step
Follow Algorithm 1 (above), taken from chapter 5 of the course booklet, step by step in Julia, to perform an LU Factorization on Matrix M.
$$ M = \begin{bmatrix}6&5&4\\12&19&14\\24&35&32\end{bmatrix}$$
We are going to do it step by step, just as if you were writing it out as a HW Drill Problem solution. Then, in the next problem, we'll program it up as an honest to goodness algorithm!

<!-- grade_id: 4, points: 0, grade: false, solution: true -->
```julia
# Do all of your work in this cell
# here, we will NOT test that the pivot is non-zero and just hope for the best!
M = [6 5 4; 12 19 14; 24 35 32]
temp=copy(M)
n,m = size(M)
k=1
C=temp[:,k]
pivot=C[k]
C=C/pivot  # Let's hope we do not divide by zero! Yikes!!! 
R=temp[k,:]' # recall, we have to do the transpose in Julia to get a row vector or we use R=temp[k:k,:]
temp = temp-C*R
@show temp
# Run this code now and make sure that temp has a zero first column and a zero first row

# Now its your turn
k=2 #now k will be 2, and we will repeat the steps again for the second column and row
#you will need to repeat this process n times, where n is the number of rows/columns in M
# your code here
throw(ErrorException())


#Now return L and U
```

<!-- grade_id: 4test, points: 1, grade: true, solution: false -->
```julia
#autograder cell
```

<img src="https://i.postimg.cc/Y0HQcjMM/lu-final.jpg" width="700">


### Problem 5. LU Factorization in Code
Implement Algorithm 1(above), taken from chapter 5 of the course booklet, in code to perform an LU Factorization on Matrix M.
$$ M = \begin{bmatrix}6&5&4\\12&19&14\\24&35&32\end{bmatrix}$$

Remark: When you first implement the algorithm using a for loop, don't worry about checking that the pivot 
is non-zero. Add that at the very end. 

Next week, we'll show you how to turn your script into a real function so that you can call
 L, U = mylu(M)

<!-- grade_id: 5, points: 0, grade: false, solution: true -->
```julia
#declare M
M = [6 5 4; 12 19 14; 24 35 32]
#initialization
n,m = size(M)
temp = copy(M)
L = Array{Float64,2}(undef, n, 0)  #~~~\# L=[] Empty matrix\\
U = Array{Float64,2}(undef, 0, n)  #~~~\# R=[] Empty matrix\\
flag = 1
#end initialization

#Your turn! Translate the rest of the algorithm to code here(use a for loop and an if statement)
# your code here
throw(ErrorException())





```

<!-- grade_id: 5test, points: 2, grade: true, solution: false -->
```julia
#autograder cell
```

#### Congratulations! You just translated an algorithm to code.  This is a key skill for a computer or data scientist.

The good news:  Julia has a built in function that will do the LU factorization for you.  Actually, there are two functions. If we use the LinearAlgebra package, we can find the L and U factors by simply typing:
<br> <br>   (L, U) = lu(M, Val(false))

<br> You can also use factorize(M), but it is not reccomended because this multiple-dispatch-function may return a different type of factorization than an LU. 

Note:  You still need to complete the algorithm above to get points.  It is important that you understand how the algorithm can be done both by hand and by Julia.  Computers exist to make your life easier, but in order for them to do that, you have to understand how they work, and how to use them.

```julia
using LinearAlgebra
#LU decomposition without permutations ~Val(false)~
(lower, upper) = lu(M, Val(false))
```

```julia
#Check that the factorization is correct
lower*upper == M
```

Note: If you leave off the Val(false) argument, lu() will use permutations, and then $L * U \ne M$


For a more in depth explanation about julia's factorization, visit:
https://julia.quantecon.org/tools_and_techniques/numerical_linear_algebra.html

### Problem 6.  Copy your algorithm from problem 5 into the cell below and run your LU Factorization on a 100 x 100 random matrix


<!-- grade_id: 6, points: 0, grade: false, solution: true -->
```julia
using Random
N=100
M=rand(N,N)
#copy your algorithm below to return L and U factors of M, a 100 by 100 matrix
# your code here
throw(ErrorException())
```

<!-- grade_id: 6test, points: 1, grade: true, solution: false -->
```julia
#autograder cell
# Check your answer
@assert isapprox(M, L*U, atol = 1E-4)
norm(M-L*U) # We have not covered norms of matrices and vectors yet in lecture, but
            # but you can think of it as a genralized notion of ``length''. 
            # Hence, the norm will be small when all of the entries of the matrix are small
            # A 100 x 100 matrix has 10,000 entries...a bit much to inspect one by one! 
```