University of Michigan - ROB 101 Computational Linear Algebra

# Homework 7: Vector Space, Subspaces and Operations
### Due 10/29 at 9 PM Eastern

#### Purpose: Use the Grahm-Schmidt process to build towards QR factorization.  Use QR factorization to create regression fits.
- Skills
    - Function building:  using functions within other functions
    - 3d surface fitting
    - Building orthogonal and orthonormal bases
- Knowledge
    - Growing undertanding of linear regression
    - Building confidence with plotting large amounts of data
    - Gearing up for project 2

### Problem 1. Building Towards QR factorization
<img src = "https://i.postimg.cc/MT0yxYcF/Grahm-schmidt.jpg" width= "800">

### a) Use the Grahm-Schmidt Process to find the orthonormal basis of the linearly independent vectors
Above is the Grahm-Scmidt Process algorithm taken from the ROB 101 booklet.
$$u1 = \begin{bmatrix}-1\\1\\-1\\1\end{bmatrix}, u2 = \begin{bmatrix}-1\\3\\-1\\3\end{bmatrix} u3 = \begin{bmatrix}1\\3\\5\\7\end{bmatrix}$$

```julia
#Step 1: declare u1, u2 and u3
u1 = 
u2 = 
u3 = 
```

<!-- grade_id: 1a, points: 0, grade: false, solution: true -->
```julia
#Step 2: Find and declare three vectors, v1, v2, and v3 that form an orthogonal basis
using LinearAlgebra 
#dot(A, B) takes the dot product 

# your code here
throw(ErrorException())
```

<!-- grade_id: 1aTest, points: 1, grade: true, solution: false -->
```julia
#autograder cell
vBasis = [v1 v2 v3]
```

### b)  Normalize your vectors v1, v2 and v3 to form an orthonormal basis
Use the norm() function of the Linear Algebra package to return the magnitude, ||v|| , of the vector.  Since normalizing is so easy in code, we will generally do it right off the bat.

<!-- grade_id: 1b, points: 0, grade: false, solution: true -->
```julia
#Normalize v1, v2 and v3 and save in v1norm, v2norm and v3norm
v1norm = v1/norm(v1)
#You do the rest
# your code here
throw(ErrorException())
```

<!-- grade_id: 1bTest, points: 1, grade: true, solution: false -->
```julia
#autograder cell
vNormBasis = [v1norm v2norm v3norm]
```

###  Because normalizing is so easy in code, we usually integrate it directly into the Gram-Schmidt Process to produce orthornormal vectors as we go. Check this out!

```julia
using LinearAlgebra 
#dot(A, B) takes the dot product 

v1 = u1
v1=v1/norm(v1)

v2 = u2 - (dot(u2,v1)/dot(v1, v1))*v1
# Thinking ahead: could you compute v2 with a for loop of length 1?
v2=v2/norm(v2)

v3 = u3 - (dot(u3,v1)/dot(v1, v1))*v1 - (dot(u3,v2)/dot(v2, v2))*v2
# Thinking ahead: could you compute v3 with a for loop of length 2? How long is the summation in (9.13) at the top of 
# the notebook?
v3=v3/norm(v3)
# Thinking ahead: what would the computation of v4 look like? And then v99? Yeah, a for loop would be good! 
#The summation in (9.13) is screaming "I am a for loop!"
println("Sweet!")
V= [v1 v2 v3]
```

### Problem 2.  Build a Grahm-Schmidt function that produces an orthonormal basis
Hint: Check out the General step shown in the Grahm-Schmidt algorithm taken from the ROB 101 booklet.  Think about how you can use code to repeat that step in order to find the orthogonal basis.  Remember to normalize as we did above!

For simplicity, your function will take the vectors stacked as the columns of a matrix U and output the orthonormal vectors as the columns of a matrix V, just as we did above.

<!-- grade_id: 2a, points: 0, grade: false, solution: true -->
```julia
#Build your function here!
function grahm_schmidt(U)
    # We assume the columns of U are linearly independent
    # because, with this assumption, norm(vk) will never be zero
    #start a blank matrix V which will hold all of the output vectors
    (rows, cols) = size(U)
    V = Array{Float64,2}(undef, rows, cols)
    #start a for loop that runs the number of times that there are columns in U
    for k in 1:cols
        #insert the general step here.  You will need another for loop to evaluate the summation
        # your code here
        throw(ErrorException())
    end
    return V
end     
```

<!-- grade_id: 2aTest1, points: 1, grade: true, solution: false -->
```julia
#test
U = [-1 -1 1; 1 3 3 ; -1 -1 5; 1 3 7]
a2Test1 = grahm_schmidt(U)
#should be the same answer that you got for v1norm, v2norm, and v3norm in Problem 1 
```

<!-- grade_id: 2aTest2, points: 1, grade: true, solution: false -->
```julia
#test 2
using Random
Random.seed!(1999)
Ubig = rand(100, 27)
sol2a = grahm_schmidt(Ubig)
```

### Problem 3.  Dependent Vectors
In Problem 2, we assumed that the given vectors were linearly independent. This will not always be the case.  Walk through the Grahm Schmidt process for vectors d1, d2 and d3, and identify where the process goes wrong.  Look back at Problem 1 part (a) for an example of how to walk through Grahm-Schmidt.
$$ d1 =  \begin{bmatrix}1\\0\end{bmatrix} , d2 = \begin{bmatrix}0\\1\end{bmatrix} , d3 = \begin{bmatrix}2\\7\end{bmatrix} $$

```julia
#Declare your d vectors
d1 =
d2 = 
d3 = 
```

```julia
#Construct your orthogonal basis
#DO NOT NORMALIZE HERE!!
v_1 = 
v_2 =
v_3 = 
```

```julia
#Normalize here
vnorm1=
vnorm2=
vnorm3=
```

<!-- grade_id: 3a, points: 1, grade: true, solution: true -->
### a) Where did we run into an error? 
Write a sentence explaining what happened, and how we can fix it so that we can still create a basis from these vectors.
YOUR ANSWER HERE

In this way, we can use the Grahm-Schmidt process to tell us if a set of vectors is linearly independent.  Essentially, we walk through the steps of GS on each vector until we get to one where $ \Vert v_k \Vert \approx 0 $.  If this is the case, then we know that $$u_k \in \operatorname{span}\{u_1, ... , u_{k-1}\} = \operatorname{span}\{v_1, ... , v_{k-1}\}$$ To say it a different way, $u_k$ caused the set $\{u_1, ... , u_{k}\}$ to be linearly dependent. Hence, we should throw it away when peforming the Gram-Schmidt process and move on to the next vector, $u_{k+1}$ and see if the set from  $\{u_1, ... , u_{k-1}, u_{k+1}\}$ is linearly independent. 

Make sense? When we encounter a dependent vector, we throw it away and move onto the next vector until we exhaust our set of vectors.

### b.  Skip the dependent vector and form an orthonormal basis from the d vectors above

<!-- grade_id: 3b, points: 0, grade: false, solution: true -->
```julia
#Remove the dependent vector, and use the other two to create the orthonormal basis
#Output the orthonormal basis in a matrix called normBasisMatrix

# your code here
throw(ErrorException())
```

<!-- grade_id: 3bTest, points: 1, grade: true, solution: false -->
```julia
#autograder cell
```

### We can also do this in a For loop, on larger sets of vectors.  Our goal below is to remove vectors that we "do not want", just as you will need to do with the "linearly dependent vectors" that Gram-Schmidt finds problematic! 

```julia
#Example where we are removing a certain "bad" vector
v1=[1; 1 ;1.0]; V=rand(3,10);V[:,1]=v1;;V[:,5]=v1;V[:,10]=v1
# Goal: remove all vectors that are equal to v1
(nrows, ncols) = size(V)
U = Array{Float64,2}(undef, nrows, 0) # no columns, yikes! U is now exists as an emptry array
for k = 1:ncols
    vk=V[:,k] #check each column vector in V
    test=norm(v1-vk) #check to see how much different the vector is from v1
    if test > 1e-8 #if it is different (difference is not = 0)
        U=[U vk] #add it to the vector U
    end
end
println("Did we remove all of the columns that we did not like? ")
U #this does not show any columns [1;1;1], so it removed what we wanted to remove
```

### Problem 4. Create a version of Gram-Schmidt that produces orthonormal vectors and eliminates linearly dependent vectors. 
Hint: Use the code from your first grahm_schmidt function, but make sure to check that you are not dividing by zero before you add another vector to the orthonormal basis.

<!-- grade_id: 4a, points: 0, grade: false, solution: true -->
```julia
#Build this new function so that it removes the dependent vectors along the way.
function grahm_schmidt_deps(U)
    (rows, cols) = size(U)
    #start a blank matrix V which will hold the orthonormal basis of vectors
    V = Array{Float64,2}(undef, rows, 0)
    #start a for loop that runs the number of times that there are columns in A
    for k in 1:cols
        #perform grahm-schmidt, check that the column is linearly independent, then if it is, add to orthonormal basis
   
        # your code here
        throw(ErrorException())
        
    end
    return V
end     
```

<!-- grade_id: 4aTest, points: 1, grade: true, solution: false -->
```julia
aTest = [0 2 4; 2 0 14]
aAns = grahm_schmidt_deps(aTest)
```

## QR factorization from the grahm-schmidt process
<img src = "https://i.postimg.cc/hGLpSYG9/QR.jpg" width = "800" >

Now that we can make an orthonormal basis from a set of linearly independent vectors, or a matrix with independent columns, we can compute the QR factorization.  We already have the Q factor!  This section taken from the ROB 101 booklet explains how to find R.

### Problem 5.  QR
### a) Using your function for the Grahm-Schmidt process, write a function that returns the QR factorization of a matrix with independent columns
Your function should return both the Q and R factors

Hint:  If you use your grahm_schmidt_deps function, the error checks for dependent vectors will be built into your QR function.

<!-- grade_id: 5a, points: 0, grade: false, solution: true -->
```julia
function qr_fact(A)
    #Calculate and declare Q and R
    # your code here
    throw(ErrorException())
    return Q, R
end
```

<!-- grade_id: 5aTest, points: 1, grade: true, solution: false -->
```julia
Q, R = qr_fact([-1 -1 1; 1 3 3 ; -1 -1 5; 1 3 7])
```

```julia
##Forward and Back substitution functions from HW04

#This is a back substitution function.  It solves for x in an equation Ux = b, where U is upper triangular.
function backwardsub(U, b)
    n = length(b)
    x = Vector{Float64}(undef, n) 
    x[n] = b[n]/U[n,n]
    for i in n-1:-1:1
        x[i]=(b[i]- (U[i,(i+1):n])' *x[(i+1):n] )./U[i,i]
    end
    return x
end
#This is a forward substitution function
function forwardsub(L, b)
    n = length(b)
    x = Vector{Float64}(undef, n); 
    x[1] = b[1]/L[1,1] 
    for i = 2:n 
        x[i]=(b[i]- (L[i,1:i-1])' *x[1:i-1] )./L[i,i] 
    end
    return x
end
```

### b) Use QR factorization to solve $Ax = b$  for a large (50x50) system of equations
<img src =  "https://i.postimg.cc/QtBJTh8B/QR-solving.jpg" width = 800>

Use the steps provided in the ROB 101 booklet, and your functions to solve.  A and b are declared for you.

```julia
#Run this to declare A and b
A = 10*rand(50, 50);
b = 10*rand(50,1);
```

```julia
#First check that the determinant of A is nonzero
det(A) != 0
```

```julia
#Then, factor A into Q and R
Q, R = 
```

```julia
#Now compute b-bar 
b_bar =
```

<!-- grade_id: 5b, points: 0, grade: false, solution: true -->
```julia
#Solve Rx = b-bar for x.  Store x in a variable x_QR
#hint: use the provided backwardsub() function
# your code here
throw(ErrorException())
```

<!-- grade_id: 5bTest, points: 1, grade: true, solution: false -->
```julia
#autograder cell
```

### c) Check your answer against an LU factorization method
Use the forwardsub() and backwardsub() functions provided.

<!-- grade_id: 5c, points: 0, grade: false, solution: true -->
```julia
#Solve for x using LU and store your answer in a variable x_LU
# your code here
throw(ErrorException())
```

<!-- grade_id: 5cTest, points: 1, grade: true, solution: false -->
```julia
#Check to ensure that your anwers are approximately equivalent
@assert isapprox(x_QR, x_LU)
```

### Problem 6.  Solve a surface regression with QR

```julia
#this cell generates the data
using Random
using Distributions
using LinearAlgebra
using Plots
#getting the data ready
z(x, y)=0.1*sin.(pi.*x./2) .- 0.1*cos.(pi.*y./3)
x = rand(Uniform(-10, 10), 100, 1)
y = rand(Uniform(-10, 10), 100, 1)
X = sort(x, dims = 1)
Y = sort(y, dims = 1)
noise = rand(Uniform(-.1, .1), 100, 1)
zdata = z(X, Y) + noise
p = scatter(X, Y, zdata)
```

**To Create the Regression model:**
1. Create a Phi($\phi$) matrix that contains all possible terms/combinations of x and y
2. Solve $zdata = \phi * \alpha$  for alpha using a QR factorization
3. The surface regression will be $zFit = \phi * \alpha$

```julia
#Creating a Phi matrix(done for you)
phi=[ones(100,1) X Y X.^2 Y.^2 X.*Y X.^3 Y.^3 (X).*(Y.^2) (Y).*(X.^2)];
```

**A reminder**
<img src =  "https://i.postimg.cc/QtBJTh8B/QR-solving.jpg" width = 600>

<!-- grade_id: 6a, points: 0, grade: false, solution: true -->
```julia
#Solve zdata = phi*alpha for alpha using your qr_fact function
#also remember to use the provided back substitution function
#store your answer for alpha in a vector called alphaStar
# your code here
throw(ErrorException())
```

<!-- grade_id: 6aTest, points: 1, grade: true, solution: false -->
```julia
#autograder cell
```

### b) Fit a surface to the data:
1. Every entry of the alphaStar array represents a term in the equation
2. Corresponding columns of phi contain the coefficients for the terms of alphaStar

Your Task:  Fill in the FOR loop so that zFit gets updated as you add terms.  Each time the for loop runs, all 100 values in zFit should be updated to include the most recent multiplication of a term in alphaStar by a column in phi.  As the loop proceeds, subsequent multiplications should be added to the previous sum.

<!-- grade_id: 6b, points: 1, grade: true, solution: true -->
```julia
#declare zFit
zFit = zeros(100)
for i in 1:10
    #Multiply the i-th term of alphaStar by the i-th column of phi, then add your result to zFit
    # your code here
    throw(ErrorException())
end

plot!(X[:,1], Y[:,1], zFit, linetype = :surface)
xlims!(-10,10)
ylims!(-10, 10)
zlims!(-1, 1)
```

### Problem 7.  Checking for existence and uniqueness of solutions to $Ax = b$ 

**Grahm-Schmidt for determining existence of solutions.**

We know that a solution to $Ax = b$ exists if $b \in \operatorname {span} \{A_{col}\} $
We also know that $b \in \operatorname {span} \{A_{col}\} $ **ONLY IF** $ \operatorname{grahm-schmidt}(A) = \operatorname {grahm-schmidt}([A, b]) $ , meaning that b is **linearly dependent** on the columns of A.

### a) Use your grahm-schmidt function to check for existence of solutions
The return statement is completed for you, so you just need to compute gsa and gsab.

<!-- grade_id: 7a, points: 0, grade: false, solution: true -->
```julia
#use your grahm_schmidt function inside this function
#Assume that A and b are size compatible
function sol_exist(A, b)
    #fill in the rest of the function
    # your code here
    throw(ErrorException())
    
    return isapprox(gsa, gsab[:,1:end-1])
end
```

<!-- grade_id: 7aTest, points: 1, grade: true, solution: false -->
```julia
#test it!
A = [1 4 6; 2 3 3; 4 5 6]
b = [1; 5; 4]
sol_exist(A, b)
```

**Null Space Computation for determining Uniqueness:**

If $Ax = b$ has a solution, it is unique only if the null space of matrix A equals the zero vector.

In julia, this looks like a Nx0 array, where N is the number of columns of A.

```julia
#An example of using the null() function to calculate null space
mat = [1 2 3; 4 5 6; 7 8 9]
nullspace(mat)
#null space was not the zero vector, therefore a system where A = mat would not have a unique solution
#You could have also discovered this because the determinant of mat is zero, so any x vector would be a solution
```

```julia
#In this example, there is a unique solution
mat2 = [1 2 3; 0 4 5; 0 0 6]
nullspace(mat2)
#notice how the output is a 3x0 Array. This is how Julia stores the 0 vector
```

### b)  Use grahm-schmidt and null space to determine if the 100x90 system below has a unique solution
The system is in the form $Jx = g$ , where J and g are declared for you below.

Hint:  Use your sol_exist function 

```julia
#This cell generates the data
using Random
Random.seed!(187)
J = 10*randn(100, 90);
g = 10*randn(100, 1);
```

<!-- grade_id: 7b, points: 0, grade: false, solution: true -->
```julia
#Return either true or false.  You must perform 2 checks
function is_sol_unique(J, g)
    #check to see if the solution exists, then if it exists, check if it is unique
    #if both conditions are met, return true.  If not, return false
    
    # your code here
    throw(ErrorException())
    
end
```

<!-- grade_id: 7bTest, points: 1, grade: true, solution: false -->
```julia
is_sol_unique(J, g)  #should ouput true
```

<!-- grade_id: 8, points: 1, grade: true, solution: true -->
### Problem 8. True or False: $x\in null(A) \iff x$ is orthogonal to each of the columns of $A^\top$
Explain your answer.  

Before answering,  either work out the answer analytically (i.e., prove or disprove the statement) OR play around with examples in Julia.  
YOUR ANSWER HERE

```julia
#Here is a space for code experimentation:

```