Unversity of Michigan ROB 101 - Computational Linear Algebra

# Homework 5.  Linear independence, Least Squares and Regression
### Due: 10/15 at 9 PM Eastern

#### Purpose:  Learn to use Julia to compute linear independence and regression models
- Skills: 
    - More practice writing functions
    - Fitting lines and curves to data
- Knowledge:
    - Understand Regression models
    - Building confidence with complex functions
    
#### Task:
Complete cells in the notebook as instructed.

## Using an LU factorization to determine whether a matrix has independent columns.
<img src = "https://i.postimg.cc/DyKPg9KX/Pro-Tip-Linear-Independence.png" width = 700>

### Example: $$A = \begin{bmatrix} 1&0&2\\2&-2&6\\3&4&2\\1&5&-3\end{bmatrix} $$

```julia
#Example:
using LinearAlgebra
#declaring the A matrix:
A = [1 0 2; 2 -2 6; 3 4 2; 1 5 -3] 
#Using Julia's LU Factorization with Row Permuation
(L, U, P) = lu(A'*A,check=false) # <---check = false avoids Julia doing something we do not need!
@show diag(U) #Show diag(U) in the Output
U
```

Notice that the last entry of diag(U) is very close to zero but not exactly zero. This is often due to roundoff error. Treating U[3,3] as being identically zero, find $$x=\begin{bmatrix}x_1\\x_2\\x_3  \end{bmatrix} \neq 0_{3 \times 1}$$
such that $A x = 0$. You will want to compute $x$ by hand! Yes, that is surprising in a jupyter notebook. Consult your 
textbook for how to determine $x$. It is not provided here. 




<!-- grade_id: cell-0507e04648baba9f, points: 0, grade: false, solution: true -->
```julia
# Given A 
A = [1 0 2; 2 -2 6; 3 4 2; 1 5 -3] 
#Use the LU factorization from above to calculate x
# Place your x-vector here (declare like x = )
# your code here
throw(ErrorException())
```

<!-- grade_id: cell-4d02da8bef32e9fa, points: 1, grade: true, solution: false -->
```julia
#autograder cell
@assert(x != [0,0,0])
@assert(isapprox(A*x, [0, 0, 0, 0], atol = 1E-4))
```

```julia
#Example
using Random
Random.seed!(20201002);  #it's a date!
A=randn(50,40); A1=A[:,1:15]; A2=A[:,2:35];A3=randn(50,5)
A=[A1 A2 A3] # building a matrix with dependent columns
#Using Julia's LU Factorization with Row Permuation
(L, U, P) = lu(A'*A,check=false) # <---check = false avoids Julia doing something we do not need!
diag(U) #Show diag(U) in the Output
```

Notice that diag(U) is too big to see all of its entries! How can we check whether it has any zero elements? 

```julia
# Run me and use me later
temp=[-6 11 -2 -1e-12 2e-3 pi]
@show minimum(temp) # finds minimum value of the vector temp
@show abs_temp=abs.(temp) 
# note the . following abs. It means to apply the absolute value
# to each entry of temp
@show minimum(abs_temp)
#
# pay extra attention here
diagU=diag(U) # diagonal of U
min_value=minimum(abs.(diagU)) 
```

### Problem 1. Build a function that returns *true* if the columns of A are Independent and *false* if they are Dependent
Use the example above to help you.  Remember, in real computations, zero is rarely identically the zero number. You have to set a threshold! How small should zero be? How can you return ``true`` vs ``false`` based on the size of a value? Hint, look through your code in juliahw3. Yes, we do expect you to keep track of what you learn. This is part of making the transition from High School to College. 

<!-- grade_id: 1, points: 0, grade: false, solution: true -->
```julia
function is_independent(A)
    #if the columns of A are independent, return true
    #if they are dependent, return false
    
    # your code here
    throw(ErrorException())
end
```

<!-- grade_id: 1aTest, points: 1, grade: true, solution: false -->
```julia
A = [1 0 2; 2 -2 6; 3 4 2; 1 5 -3] 
is_independent(A)
```

<!-- grade_id: 1bTest, points: 1, grade: true, solution: false -->
```julia
aInd = [7   2   1   2   9;1   1   6   8   2;1   10   5   10   7;8   6   8   2   3;8   4   5   5   10]
is_independent(aInd)
```

### Problem 2.  Existence and Uniqueness of Solutions.
<img src = "https://i.postimg.cc/3wHhJJMw/Semi-Pro-Tip-Number-Linearly-Independent-Vectors.png" width = 600>


### a) Consider the system of linear equations give below,
$$\begin{equation}
\label{eq:LinearCombinationViaLU}
\underbrace{\left[\begin{array}{rrr}
 1 & -2 &  1 \\
 2 &  0 & 0 \\
 3 &  -3 &  1 \\
 4 &  1 &  0 \\
 5 &  -4 &  1
 \end{array}\right]}_{A} \underbrace{\left[\begin{array}{c}
x_1\\ x_2\\x_3 \end{array}\right]}_{x} =  \underbrace{\left[\begin{array}{r}
9\\ 4\\ 16 \\5 \\ 23 \end{array}\right]}_{b}.
\end{equation}$$
### Based on the method described above, perform the relevant LU Factorizations associated with the matrices $A$ and $[A ~~b]$, respectively. Store the diagonals of the two upper triangular matrices in ``diagU`` and ``diagUe``. 

## Warning: The autograder is looking for their diagonals and not $U$ or $U_e$! 


<!-- grade_id: 2a, points: 0, grade: false, solution: true -->
```julia
# a) part 1
A = [1 -2 1; 2 0 0; 3 -3 1; 4 1 0; 5 -4 1]
b = [9; 4; 16; 5; 23]
# Perform relevant LU Factorizations associated with the matrices A and [A b] respectively.
# Store the diagonals of the two upper triangular matrices in variables called:
#              diagU and diagUe
# your code here
throw(ErrorException())
```

<!-- grade_id: 2aTest, points: 1, grade: true, solution: false -->
```julia
#autograder cell
```

### b) Based on your analysis of the diagonals, state whether  $Ax = b$ has a solution, or if it cannot be determined by only looking at the diagonals. 
### Provide your answer by uncommenting one of the following lines.

To "uncomment" a line, just delete the # character in the front of the line.  This will change that line from a comment to actual, executable code.

```julia
# Ans=false
# Ans=true
# Ans=Array{Int}(undef, (1, 1)) 
```

<!-- grade_id: cell-c7c00d831328e560, points: 1, grade: true, solution: false -->
```julia
#autograder cell
```

### Problem 3.  The Norm of a Vector
<img src = "https://i.postimg.cc/ZnngJ7zk/normDef.jpg" width = 700>

### a)  Calculate the Norm of vector v1 $$v1 = \begin{bmatrix}-1\\4\\7\\3\\6\\2\end{bmatrix}$$

Hint:  Use the definition above, taken from the ROB 101 booklet.

<!-- grade_id: 3a, points: 0, grade: false, solution: true -->
```julia
#vector v1 is declared for you
v1 = [-1; 4; 7; 3; 6; 2]
#store the norm into a variable called normv1
# your code here
throw(ErrorException())
```

<!-- grade_id: 3aTest, points: 1, grade: true, solution: false -->
```julia
#autograder cell
```

### b.  Given a vector of an unknown length, write a function that calculates the norm of the vector


<!-- grade_id: 3b, points: 0, grade: false, solution: true -->
```julia
#Fill in the necessary elements of the function
function my_norm(v)
    #initializing a sum variable to be zero
    sumSquares = 0
    #write a For loop that adds up the squares of all the elements of v
    # your code here
    throw(ErrorException())
    #Now take the root of the sum of all the squares, and return the value from the function
    # your code here
    throw(ErrorException())
end
```

<!-- grade_id: 3bTest, points: 1, grade: true, solution: false -->
```julia
#check that your function returns the same value as you got earlier for normv1
my_norm(v1) == normv1
```

**The secret:  Julia has a Built-in Norm function in the linear algebra package, called norm().  You can use it the exact same way as your my-norm() function.**

Note:  You still must complete the function my_norm() as indicated to receive points.  This is not because we want you to do unnecessary busy work, it is because we think it is important that you understand what Julia is doing behind the scenes when you use an included function.

### Problem 4. Least Squares Solutions to linear equations
<img src = "https://i.postimg.cc/7hMvLBKm/LSQ.jpg" width = 700>

### a)  Find a least squares solution to $Ax = b$, given
$$\begin{equation}
\label{eq:LeastSquareSolExample}
\underbrace{\left[\begin{array}{rrr}
 1.0 & 1.0 \\
 2.0 & 1.0 \\
 4.0 & 1.0 \\
 5.0 & 1.0 \\
 7.0  & 1.0
 \end{array}\right]}_{A} \underbrace{\left[\begin{array}{c}
x_1 \\ x_2  \end{array}\right]}_{x} =  \underbrace{\left[\begin{array}{r}
4 \\  8 \\ 10 \\ 12 \\ 18 \end{array}\right]}_{b},
\end{equation}$$
A reminder about LU with permutations for solving $Ax = b$:
<img src= "https://i.postimg.cc/D0NFCzXc/LU-P.jpg" width = 700>

```julia
#Declare A and b
A = [1 1; 2 1 ; 2 1; 5 1; 7 1]
b = [4 8 10 12 18]'
```

```julia
#Compute A' * A and store in a variable aTransA
aTransA = A'*A
```

```julia
#Compute A' * b and store in a variable aTransB
aTransB = A'*b
```

```julia
# Run me, don't change me. I am providing some nice functions for you.
##Forward and Back substitution functions from HW04

#This is a back substitution function.  It solves for x in an equation Ux = b, where U is upper triangular.
# The function assumes the diagonal elements are non-zero. You can change that if you wish!
function backwardsub(U, b)
    n = length(b)
    x = Vector{Float64}(undef, n) 
    x[n] = b[n]/U[n,n]
    for i in n-1:-1:1
        x[i]=(b[i]- (U[i,(i+1):n])' *x[(i+1):n] )./U[i,i]
    end
    return x
end
#This is a forward substitution function. It solves for x in an equation Lx = b, where L is lower triangular.
# The function assumes the diagonal elements are non-zero. You can change that if you wish!
function forwardsub(L, b)
    n = length(b)
    x = Vector{Float64}(undef, n); 
    x[1] = b[1]/L[1,1] 
    for i = 2:n 
        x[i]=(b[i]- (L[i,1:i-1])' *x[1:i-1] )./L[i,i] 
    end
    return x
end
```

<!-- grade_id: 4a, points: 0, grade: false, solution: true -->
```julia
#Compute x by solving    aTransA * x = aTransB
#store x in a variable called xStar

#use an LU factorization with permutations, and the forwardsub() and backwardsub() functions provided
# your code here
throw(ErrorException())

#Hint:  Make sure you are solving A' * A x = A' b
```

<!-- grade_id: 4aTest, points: 1, grade: true, solution: false -->
```julia
#autograder cell
```

### b.  Compute the error and determine if we found an exact solution
Remember, e* = Ax* - b

```julia
#Compute e* 
eStar = A*xStar-b
```

```julia
#Compute ||e*|| 
normE = norm(xStar)
```

```julia
#Compute ||e*||^2 
normEsquared = normE^2
```

<!-- grade_id: 4c, points: 1, grade: true, solution: true -->
### c.  Does this system have an exact solution?  Explain why or why not.

YOUR ANSWER HERE

### Problem 5.  Linear fitting with Least Squares
Recall that a linear regression can be written as the linear equations:
$$y_i = m x_i +b = \begin{bmatrix} x_i & 1 \end{bmatrix} \begin{bmatrix} m \\ b \end{bmatrix},  ~~1 \le i \le N,$$ 
Where N is the number of data points.

In Matrix Form:
$$\begin{equation}
    \label{eq:FirstRegressionModel02}
\underbrace{\begin{bmatrix} y_1 \\ y_2 \\ \vdots \\y_N \end{bmatrix}}_{Y} = \underbrace{\left[\begin{array}{cc}
    x_1 & 1 \\
    x_2  & 1 \\
    \vdots & 1 \\
    x_N & 1
\end{array}  \right]}_{\Phi} \cdot  \underbrace{\begin{bmatrix} m \\ b \end{bmatrix}}_{\alpha},
\end{equation}$$
Where $Y$ is the vector of $y$-data, $\Phi$ is called the regressor matrix and $\alpha$ is the vector of unknown coefficients that parameterize the  model.

```julia
#This cell generates the data
using Random
Random.seed!(1234)
noise = randn(300, 1)
xvals = collect(0:.1:29.9)
yvals = 1/2*xvals + noise
using Plots
scatter(xvals, yvals)
```

### a)  Create the Y and $\Phi$ Matrices for the data.
Look back at the problem 5 cell for help identifying which data goes in which matrix.

```julia
#Declare Y and Phi below
Y = yvals
Phi = [xvals ones(300,1)]
```

### b) Find the least squares solution to $Y =  \Phi  \alpha $  
Hint:  Follow the steps you followed in Problem 4 to use the LU factorization, except we have $\Phi$ instead of $A$, $\alpha$ instead of $x$, and $Y$ instead of $b$

<!-- grade_id: 5b, points: 0, grade: false, solution: true -->
```julia
#compute alpha* and return it in a variable called alphaStar
# your code here
throw(ErrorException())
```

<!-- grade_id: 5bTest, points: 1, grade: true, solution: false -->
```julia
#autograder cell
```

### c)  Plot the least squares solution on the same plot as the scatter.  
Reminder: $\alpha = \begin{bmatrix} m \\ b \end{bmatrix}$ and $Y = mx + b$

<!-- grade_id: 5c, points: 1, grade: true, solution: true -->
```julia
#Plot the regression solution on the same plot.  Hint: Use plot!()
# your code here
throw(ErrorException())

#hint:  if you are getting a MethodError: no method matching +(::Array{Float64,1}, ::Float64),
#check that you are not trying to add a scalar to an array. 
```

### d)  Compute the total squared error
$$E_{tot} = \sum_{i=1}^{5} (e_i)^2 = e^\top e = ||e||^2= ||Y-\Phi \alpha ||^2.$$
Use this equation to give an objective measure of how accurate your linear fit was.

<!-- grade_id: 5d, points: 0, grade: false, solution: true -->
```julia
#Use whatever method you would like, but store the value in a variable totSquareError
# your code here
throw(ErrorException())
```

<!-- grade_id: 5dTest, points: 1, grade: true, solution: false -->
```julia
#autograder cell
```

### Problem 6.  Quadratic Fitting
You can also fit a regression to a non-linear set of data.  It is still a "linear regression" even if the data itself is not a linear pattern.

Writing out the quadratic regression equation $y_i = c_0 + c_1 x_i + c_2 x_i^2$ $,i=1,\cdots,N$ in matrix form yields

$$\underbrace{\left[ \begin{array}{c} y_1 \\ y_2 \\ \vdots \\ y_N \end{array} \right]}_{Y} =
 \underbrace{\left[ \begin{array}{ccc} 1 & x_1 & (x_1)^2 \\ 1 & x_2 & (x_2)^2 \\ \vdots & \vdots \\ 1 & x_N & (x_N)^2 \end{array} \right]}_{\Phi}
 \underbrace{\left[ \begin{array}{c} c_0 \\ c_1 \\ c_2 \end{array} \right] }_{\alpha},
$$

```julia
#This cell generates the data
using Distributions
using Plots
Random.seed!(56789)
f(x) = 3*x+2*x.^2 
N=300
x=rand(Uniform(-2,2),300,1)
x=sort(x,dims = 1)
noise=rand(Uniform(-0.5,0.5),300,1)
ydata=f(x) + noise
scatter(x,ydata)
```

### a)  Create the $Y$ and $\Phi$ Matrices 

```julia
#Declare your matrices below
phi = [ones(N,1) x x.^2]
Y = ydata
```

### b) Find the least squares solution to $Y =  \Phi  \alpha $ 

<!-- grade_id: 6b, points: 0, grade: false, solution: true -->
```julia
#compute alpha* and return it in a variable called alphaStar
# your code here
throw(ErrorException())
```

<!-- grade_id: 6bTest, points: 1, grade: true, solution: false -->
```julia
#autograder cell
```

### c)  Plot the least squares regression on the same plot as the scatter
Remember: $ \alpha = \left[ \begin{array}{c} c_0 \\ c_1 \\ c_2 \end{array} \right]$, and the fit is $y_i = c_0 + c_1 x_i + c_2 x_i^2$

<!-- grade_id: 6c, points: 1, grade: true, solution: true -->
```julia
#Declare y_fit and plot it on the same graph with plot!()
# your code here
throw(ErrorException())
```

### d)  Compute the total squared error

<!-- grade_id: 6d, points: 0, grade: false, solution: true -->
```julia
#Use whatever method you would like, but store the value in a variable totSquareError
# your code here
throw(ErrorException())
```

<!-- grade_id: 6dTest, points: 1, grade: true, solution: false -->
```julia
#autograder cell
```