University of Michigan - ROB 101 Computational Linear Algebra

# Homework 2: Vectors, Matrices and Linear Systems
### Due:  9/17 at 9 PM eastern

#### Purpose:  Become comfortable using vectors and matrices in code 
- Skills: 
    - Declaring vectors and matrices
    - Indexing into vectors/matrices
    - Writing systems of linear equations in matrix form
- Knowledge:
    - Matrix manipulation (central to linear algebra)

#### Task: For this assignment, fill in the code blocks as directed to achieve the desired output.

### Problem 1.  Indexing (1+1+1 points)
### a) return the biggest number in the matrix bigNumbers with indexing

<!-- grade_id: 4a, points: 0, grade: false, solution: true -->
```julia
bigNumbers = [1 34 6 56 7 234 6 5; 3 43635 2 35 6 234 6 23; 235 56 2 46 56 3 234 6; 87 23 8 23 9 47 0 8]

#store the biggest number in a variable called biggestNum
# biggestNum = ??
# your code here
throw(ErrorException())

#return the biggest number 
println("Biggest: ", biggestNum)
```

<!-- grade_id: 4aTest, points: 1, grade: true, solution: false -->
```julia
#autograder cell
#this cell will not output anything but make sure to run it so your answer is checked
```

### b) return the column containing the biggest number as a column vector called bigColumn

Hint:  In indexing, the colon : operator in place of a row or column indicates that all the rows/columns will be selected

<!-- grade_id: 4b, points: 0, grade: false, solution: true -->
```julia
#Use indexing to return bigColumn
# bigColumn = ??
# your code here
throw(ErrorException())
```

<!-- grade_id: 4bTest, points: 1, grade: true, solution: false -->
```julia
#autograder cell
```

### c) return the row containing the biggest number as a row vector

Hint:  You will probably need to transpose this vector.  Julia will automatically return a column vector, even though you are selecting a row, so you will need to transpose to convert it to a row vector with a single quote( ' ) after the vector statement.


```julia
#Example of transposing
# run me don't change me, just an example of transposing
@show Z = [1, 2, 3, 4]  #Z is a column vector
@show Z'
#Z' is a row vector
```

<!-- grade_id: 4c, points: 0, grade: false, solution: true -->
```julia
#Return the row as bigRow
# bigRow = ??
# your code here
throw(ErrorException())
```

<!-- grade_id: 4cTest, points: 1, grade: true, solution: false -->
```julia
#autograder cell
```

#### Elements of an Array vs Scalars
In Julia, there is a difference in type between a 1x1 matrix and a number.  Silly and annoying, but easy to fix.  By using indexing, we can take the number out of the matrix.  We simply use matrix[1]

### Problem 2. (1 point) 
Extract the number from the 1x1 matrix, and put it in a variable called numberInside

<!-- grade_id: 5, points: 0, grade: false, solution: true -->
```julia
tinyMatrix = [7.01]
#output the number inside tinyMatrix into a variable of type Float64 called numberInside
# numberInside = ??
# your code here
throw(ErrorException())
```

<!-- grade_id: 5test, points: 1, grade: true, solution: false -->
```julia
#autograder cell
@assert numberInside == 7.01
```

### Problem 3. (1+1 points)
This 3 by 3 System of linear equations is from Problem 3 of the written homework.

$$ 4x_1  = 8 $$
$$  2x_1 - 10x_2 = -2 $$
$$ x_1 + 2x_2 - x_3 = 4 $$

### a) Create matrices A and b so that Ax = b represents this 3 by 3 system.  $x = \left[x1; x2; x3\right]$

<!-- grade_id: 6a, points: 0, grade: false, solution: true -->
```julia
#Declare A and b
# A = ??
# b = ??
# your code here
throw(ErrorException())
```

<!-- grade_id: 6aTest, points: 1, grade: true, solution: false -->
```julia
#autograder cell
```

Now that you found the matrices that define this system, the algorithm below shows how we can use forward substitution in code to solve for the x vector.  
### b) Display the vector x in the output window with whatever method you prefer

<!-- grade_id: cell-6841f35a7a410414, points: 1, grade: true, solution: true -->
```julia
x = [0.0; 0.0; 0.0] #Initialize an empty x vector

#This method does not scale to large problems, but it is super clear
x[1]=b[1]/A[1,1]
x[2]=(b[2]-A[2,1]*x[1])/A[2,2]
x[3]=(b[3]-A[3,1]*x[1]-A[3,2]*x[2])/A[3,3]
;
#add a line here that will display the vector x
# your code here
throw(ErrorException())
```

### Problem 4. (3 points)
**Now, modify the algorithm above to work for a 5 by 5 system.**

$$ 4x_1  = 8 $$
$$  2x_1 - 10x_2 = -2 $$
$$ x_1 + 2x_2 - x_3 = 4 $$
$$ 4x_1 - x_2 - 3x_3 + x_4 = 7 $$
$$ -2x_1 + 7x_2 + x_3 - 2x_4 - 6x_5 = 12 $$


<!-- grade_id: 7a, points: 0, grade: false, solution: true -->
```julia
#first, declare your A and b
# A = ??
# b = ??
# your code here
throw(ErrorException())
```

<!-- grade_id: 7aTest, points: 1, grade: true, solution: false -->
```julia
#autograder cell
```

<!-- grade_id: cell-99064965fc2bda65, points: 0, grade: false, solution: true -->
```julia
x = Vector{Float64}(undef, 5)  #this time, x will be a vector of 5 entries(this is another way to start a vector)

#now, modify the algorithm to work on this 5 by 5 system by adding two more equations.
x[1]=b[1]/A[1,1]
x[2]=(b[2]-A[2,1]*x[1])/A[2,2]
x[3]=(b[3]-A[3,1]*x[1]-A[3,2]*x[2])/A[3,3]
#add your two equations below as x[4] and x[5]
# your code here
throw(ErrorException())
```

<!-- grade_id: cell-1baf8eae874023ec, points: 2, grade: true, solution: false -->
```julia
#autograder cell
```

**That was super tedious, right? If not, then we failed. You were supposed to go, what are they thinking, there must be a better way to go about this!! Yes, there is.** 

### A Better Algorithm!

But first, let's generate some data to set you up for success. Run the next cell. There is nothing to report. 

```julia
# Execute this line to define A, b and x before you move on 
x = [0.0; 0.0; 0.0] #Initialize an empty x vector
A=[4.0 2.1 1.0; -7.03 1.0/3.0 π; 2.7 -18.1 -11.11]
b=[15.0 13.6 -21.2]'
```

**The problem with the "steps" in the previous cell is that each line of calculation is different from the others. That makes it hard to "automate", which is what programming is all about right? We humans set up the program and we let the computer automtically grind through the computations?**

Let's focus on the very first line of the forward substitution method
$$ x[1]=b[1]/A[1,1] $$

Suppose that we set x to be a zero column vector with three elements 
$$ x=[0.0;0.0;0.0] $$ 

You know that $ A[1,:]'$ is the first ROW of the A matrix

What will be the result of 
$ A[1,:] '*x $ ?

Because x is the zero vector we get zero. We weren't trying to fool you, just making you admit the obvious.

Hence if we could write something like
$$ x[1]=(b[1] -  A[1,:]' *x )/A[1,1] $$
we woudl get the same result as
$$ x[1]=b[1]/A[1,1] $$
which is the first step of forward substitution. Let's note that we have updated the value of $ x[1] $, while $ x[2] $ and $ x[3] $ remain zero. Why is this important? 

Well, ask yourself, what will be the result of the multiplication  of the Second row of A times x?
$$ A[2,:] '*x $$   
                                                
once again, knowing that $ x[2]=0 $ and $ x[3]=0 $ ?

Because only $x[1]$ is potentially non-zero, we have $A[2,1]*x[1]$

Hence if we could write something like
$$x[2]=(b[2] -  A[2,:]' *x )/A[2,2]$$
we would get the same result as
$$x[2]=(b[2]-A[2,1]*x[1])/A[2,2]$$

Hopefully you see the pattern, if we could write something like 
$$x[3]=(b[3] -  A[3,:]' *x )/A[3,3]$$
we'd get the same result as 
$$x[3]=(b[3]-A[3,1]*x[1]-A[3,2]*x[2])/A[3,3]$$

Now, all of our lines of code look the same. Therefore, we can automate them! 

There is one catch, which maybe you already noted: in Julia a row vector times a column vector (of the same lengths) gives a 1x1 matrix, which is not the same thing as a scalar (that is, a number). So, we have to extract out (or slice out) that number via: 
$$ (A[k,:])' * x)[1]$$

You may need to go back to the top and look at Problem 5 and tinyMatrix.

In the next cell, you can run the two sets of equations and verify that we get the same answers. 

```julia
# New way where there is a pattern
x=[0.0;0.0;0.0]
x[1]=(b[1]-(A[1,:]' *x)[1])/A[1,1]
x[2]=(b[2]-(A[2,:]' *x)[1])/A[2,2]
x[3]=(b[3]-(A[3,:]' *x)[1])/A[3,3]
@show x
# Old way where each line is different
xx=[0.0;0.0;0.0]
xx[1]=b[1]/A[1,1]
xx[2]=(b[2]-A[2,1]*xx[1])/A[2,2]
xx[3]=(b[3]-A[3,1]*xx[1]-A[3,2]*xx[2])/A[3,3]
@show xx
```

We're almost at a crucial point of having our first real program. **Almost!** 

We'll let you grind a bit more! Ouch! 

### Problem 5. (1 point) 
For a system of equations Ax = b that is bigger, for example, 10x10, it would make more sense to have some sort of repeating code, than to type 10 equations. We've seen that our magic common equation has the form

$$x[k]=(b[k]-(A[k,:]' * x)[1])/A[k,k]$$ 

For this problem, we will give you a large A and b. You assume that Ax = b, and that x is a column vector with 10 entries. We'll have you MANUALLY repeat a step 10 times! (Hey, it's better than typing 10 equations!)

```julia
#Run this cell to generate matrices A and b, and initialize x to be the correct size.  DO NOT MODIFY IT!
#Once it is done running, you should see a 10x10 lower triangular matrix as the output

using Random
Random.seed!(1234)
N=10; A=rand(N,N); #Builds a 10 x 10 matrix filled with randon elements, except, we fixed the 
# "seed" so that all of you will generate the same matrix

# Make A lower triangular with one for loop and indexing
# For now, pretend that you have not seen the command "for" and its friend "end"
for k=1:N
    A[k,k+1:N]=0.0* A[k,k+1:N]; #Not important to understand this line at this TIME!
end 
b=rand(N,1); # a random column vector of size 10 x 1. Not bad.
x=0*b; # Sets x to zero and has the same size as b
@show size(x)
A
```

Now, to find all the values of x, we can define an equation that looks like our  algorithm above, 
$$ x[k]=(b[k]-(A[k,:]' *x)[1])/A[k,k]  $$  
and run it 10 times.

```julia
#first we need to start a counter variable to count the times that we have run the loop and so that we can index into the
# x vector
k = 0
```

```julia
#you will need to run this block of code for EVERY ROW of the matrix A (10 times).  
#Each output will correspond to an entry in the x vector.  Enter them in the block below for grading.
k = k + 1
x[k] = (b[k]-((A[k , :]') *x)[1])/A[k,k]
println(k)
println(x[k])

#hint: If you get a BoundsError, re-run the cell above to reset k to zero. 
#you will also need to re-run the initialization cell
```

```julia
#Enter your results here as x1 through x10:
x1 = 
x2 = 
x3 =
x4 =  
x5 = 
x6 = 
x7 = 
x8 = 
x9 = 
x10 =  
;
```

<!-- grade_id: 8Test, points: 1, grade: true, solution: false -->
```julia
#autograder cell
```

#### Using the inv() function of the Linear Algebra package.  
This function takes in the matrix A and the vector b, and returns the solution vector x.  No, this is not cheating, it is checking your work!  The cell below should return the same vector as your solution to the cell above.

```julia
#Your results should match this
using LinearAlgebra
inv(A)*b
```

#### What if, instead of running that cell over and over manually, we used a control flow operator to do it for us?  For loops are perfect for running a certain number of times, and they even keep track of the counter variable for us.

### Problem 6. (3 points)
Find the solution to a 100 by 100 system! 

```julia
# This generates the matrices A and b for a 100x100 system.  DO NOT MODIFY THIS CELL!

Random.seed!(4321)
N=100; A=rand(N,N);
# Make A lower triangular with one for loop and indexing

# This time, pay attention to "for" and its friend "end". Everything in between is repeated for k=1 all the way 
# up to k=N. How convenient is that? (Yes, we know, when we first learned this we were bowled over too!)

for k=1:N  
    A[k,k+1:N]=0.0* A[k,k+1:N]; #Not important to understand this line at this TIME!
end 
b=rand(N,1);
x=0*b;  #makes an x vector the same size as the b vector, but full of zeros

```

<!-- grade_id: 9, points: 0, grade: false, solution: true -->
```julia
for k=1:N
    #fill in a line that calculates x[k]  Hint:  Look at the cell you repeated earlier
    # your code here
    throw(ErrorException())
end
print(x[1:10, 1]) #show the first 10 entries of x
```

#### Check your work with the inv() function.

<!-- grade_id: 9test, points: 3, grade: true, solution: false -->
```julia
#the inv() function
ansComp = inv(A)*b
print(x[1:10, 1])
@assert isapprox(x, ansComp) #this function will tell us if the computer got the same values with inv() as we did with the algorithm
#if it returns "true" then your algorithm worked! 
#Otherwise, try again, and make sure to re-run the cell that builds the matrices to get a fresh start.
```

Cheatsheets for Julia Vector and Matrix syntax
* https://cheatsheets.quantecon.org/julia-cheatsheet.html#vectors-and-matrices
* https://www.softcover.io/read/7b8eb7d0/juliabook/basics

Documentation for the LinearAlgebra Package
* https://docs.julialang.org/en/v1/stdlib/LinearAlgebra/